## Instructions (important, please read)

1. Type your full name where it says NAME, below

2. Put your solution into the cell that says:
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**(Delete these contents when inputing your solution)**

2. Do not move, add, or delete any cells

3. Make sure your variable names and strings **match the instructions exactly**, (e.g. variable_1 vs Variable_1)

3. If you are not using Colab, please copy and paste your **final answers into Colab**

4. Before you submit, make sure everything runs as expected: 
    * Input your final answers  
    * Go to "Runtime" $\Rightarrow$ "Restart and Run All" 
    * Make sure you don't see any errors

5. Go to "File" $\Rightarrow$ "Download .ipynb"

6. Upload the .ipynb to Canvas (do not change the file name)

**Important Note:** There is an empty cell below your solution, **do not move or delete it** (it contains hidden tests that assign points to your code, deleting it will result in 0 points).

```python
NAME = "Ahmad Rafi Faqiri"
```

---

# k-Nearest Neighbors

**Problem:** Create a .py file called 'vec_alg_helper_fxs.py' that contains all the functions defined in the previous homework assignment. 


From the vec_alg_helper_fxs.py file, use an `import` statement to import `vec_distance`.

<!-- grade_id: cell-abda05f832b82276, points: 0, grade: false, solution: true -->
```python
from vec_alg_helper_fxs import vec_distance

```

<!-- grade_id: cell-3c12843dc1a17b26, points: 5, grade: true, solution: false -->
```python

```

**Problem**

Write a function ``majority`` that returns a majority vote, given a list of labels (strings). 

For example, if the argument of the function is ``name_list = ['John', 'Anna', 'Anna', 'Tom', 'Marie']``, then the call ``majority(name_list)`` should return ``Anna`` since it appears the greatest number of times in the argument.

(Don't worry about this now, but assume label order matters and they are pre-sorted from closest to farthest) 

**Function name:** `majority()`

**Input arguments:** 

* the list of labels: `List[str]`

**Output:** 

* the majority vote (most common label): `str`

**Requirements:**

* In the case that there is a tie (two or more winners), remove the farthest (the last element of the list) and re-run the function.
* Use a recursive function to accomplish this. (This means the `return` value calls the function again: https://realpython.com/python-thinking-recursively/)

**Hints:**

* Use `Counter` (`from collections import Counter`)
* `Counter.most_common(1)` can be used to return the most common count from a list
* Use list comprehension to see if there are any ties



<!-- grade_id: cell-ad59d584214b14d6, points: 0, grade: false, solution: true -->
```python
from collections import Counter
def majority(labels):
  label_counts=Counter(labels)
  vote=label_counts.most_common()
  if len(vote)==1:
    return vote[0][0]
  if vote[0][1]!=vote[1][1]:
    return vote[0][0]
  if vote[0][1]==vote[1][1]:
    labels.pop()
    maj=majority(labels)
    return maj



   
```

```python
maj=majority(['a','b','b','a'])
```

```python
print(maj)
```

Output:
```
b
```

<!-- grade_id: cell-0ce3da6f315b9536, points: 5, grade: true, solution: false -->
```python
"""Here's one of the tests I'll run on your code:"""
assert majority(['a', 'a', 'b','c','a','b']) == 'a'
```

**Problem**

Write a class that represents a labeled data point. It should have two attributes: 

* a point (represented as a list)
* a label (represented as a string)

For example, once your class exists, you should be able to do this:

`a = LabeledPt(point=[0,0], label='my_point')`

And be able to access its attributes like this:

`a.point = [0,0]` and `a.label = 'my_point'`

**Class name:** `LabeledPt()`

**Class attributes:**

* `point`: list
* `label`: string

**Hint:**

Remember to use the ``__init__`` method as ```def __init__(self, point, label):```

<!-- grade_id: cell-fd725bed29e1967c, points: 0, grade: false, solution: true -->
```python
class LabeledPt():
    
    def __init__(self, point, label):
      self.point=point
      self.label=label
    def point(self):
      pt=self.point
      return pt

    def label(self):
      lb=self.label
      return lb


```

```python
a = LabeledPt(point=[0,0], label='my_point')
```

```python
a.point
```

Output:
```
[0, 0]
```

<!-- grade_id: cell-5b6f4ca38b654321, points: 5, grade: true, solution: false -->
```python
a.label
```

Output:
```
'my_point'
```

**Problem**

Write a KNN classifier function `knn_classifier()`:

Given a list of labeled points (`labeled_points`) and a new point (`new_point`) to classify, the function should return the majority vote of the labels of the `k` nearest neighbors to `new_point`.

By "nearest" we mean the shortest `vec_distance()` between points

For example, if we have two labeled points in the list `labeled_points = [LabeledPt(point=[0,0], label='a'), LabeledPt(point=[1,1], label='b')]`, then calling the `knn_classifier` function to return the label for the new point `new_point=[0.49, 0.49]` based on the `k=1` nearest neighbor should output `a`:

```python
knn_classifier(1, 
                 [LabeledPt(point=[0,0], label='a'), 
                  LabeledPt(point=[1,1], label='b')], 
                 new_point = (0.49,0.49)
              ) == 'a'
```

**Function name:** `knn_classifier()`

**Input arguments:** 

* `k`: the number of neighbors
* `labeled_points`: a list of labeled points
* `new_point`: the point you're trying to classify

**Output:** 

* the predicted class label of the new point

**Hints:**

* Use `majority()`
* Use the built-in Python function `sorted()` with a custom `Key` function involving `lambda` and the ``vec_distance()`` method imported at the beginning of the homework assignment

```python
knn_classifier(1, 
                 [LabeledPt(point=[0,0], label='a'), 
                  LabeledPt(point=[1,1], label='b')], 
                 new_point = (0.49,0.49)
              ) == 'a'
```

Output:
```
[{'label': 'a', 'distance': 0.6929646455628166}]
[{'label': 'a', 'distance': 0.6929646455628166}, {'label': 'b', 'distance': 0.7212489168102785}]
a
a
True
```

<!-- grade_id: cell-ab90c43fc1344483, points: 0, grade: false, solution: true -->
```python
def knn_classifier(k,labeled_points,new_point):
  L=[]
  D={}
  V=[]
  for lbpt in labeled_points:
    a = lbpt
    


    dist=vec_distance(a.point, new_point)
    # D['label']=a.label
    # D['distance']=dist
    # print(D)
    L.append({'label':a.label,'distance':dist})
    print(L)
  sor = sorted(L, key=lambda s: s['distance'])
  # print(sor)
  for i in range(k):
    votes=sor[i]['label']
    print(votes)
    V.append(votes)
  pred_label=majority(V)
  print(pred_label)
  return pred_label



   

```

<!-- grade_id: cell-8c889fa4c4cc150a, points: 10, grade: true, solution: false -->
```python

```

**Run code below to load the iris dataset**

This is one of the best known classification problems in machine learning, and is frequently used to "benchmark" machine learning algorithms.

The goal is to classify the *flower species* (represented as `strings`):

* Setosa
* Versicolor
* Virginica

As a function of *four variables* (represented as `floats`):

* Sepal length
* Sepal width
* Petal width
* Petal lenth

You can read about it here:

https://en.wikipedia.org/wiki/Iris_flower_data_set


```python
import csv

with open('iris.dat') as f:
    read_file = csv.reader(f)
    loaded_iris_arr = [row for row in read_file]
```

**Background**

We have just loaded the Iris data as a list of lists. One row (i.e. one sample) looks like this:

```python

loaded_iris_arr[0]

['5.1', '3.5', '1.4', '0.2', 'Iris-setosa']
```

Where `5.1`, `3.5`, `1.4`, `0.2` are strings representing sepal length, sepal width, petal length, and petal width respectively (all in centimeters). `'Iris-setosa'` is a string representing the species.

**Problem**

Write a function, `gen_LabeledPt()` that converts each row of the *Iris* dataset to a `LabeledPt`. 

For example, once your function is created you should be able to do this:

```python

test =  gen_LabeledPt(loaded_iris_arr[0])
print(test.point)
print(test.label)

[5.1, 3.5, 1.4, 0.2]
setosa
```

**Function name:**

* `gen_LabeledPt()`

**Input arguments:** 

* a row from `loaded_iris_arr`. For example `loaded_iris_arr[0]`

**Output:** 

* The function should return an instance of class `LabeledPt`

**Requirements:**

* Remove `'Iris-'` from the species string, e.g. if the class is *Iris-setosa*, the `.label` should be ```setosa```

**Hints:**

* Use the ``split()`` function to remove `'Iris-'`
* `measurement_values` can be created using a list comprehension

```python
h=loaded_iris_arr[100]
```

```python
h[4]
```

Output:
```
'Iris-virginica'
```

```python
h[4].split('-')[1]
```

Output:
```
'virginica'
```

```python
test =  gen_LabeledPt(loaded_iris_arr[0])
print(test.point)
print(test.label)

```

Output:
```
['5.1', '3.5', '1.4', '0.2']
setosa
```

<!-- grade_id: cell-9741ecf2cfcf2350, points: 0, grade: false, solution: true -->
```python
def gen_LabeledPt(row):
    measurement_values=row[0:4]
    label=row[4]
    point_label=label.split('-')[1]
    return LabeledPt(measurement_values, point_label)
```

<!-- grade_id: cell-cbed73f84c3f77ad, points: 10, grade: true, solution: false -->
```python

```

**Problem**

Use list comprehension and your `gen_LabeledPt` function to convert *all* the rows of the `labeled_iris_points` to `LabeledPts`

**Output:** 

* The output should be a list of 150 LabeledPts called ```labeled_iris_points```



```python
gen_LabeledPt(loaded_iris_arr[0])
```

<!-- grade_id: cell-eddaaefad6a56a66, points: 0, grade: false, solution: true -->
```python
labeled_iris_points=[gen_LabeledPt(i) for i in loaded_iris_arr[0:150] ]
```

```python
labeled_iris_points[0].label
```

Output:
```
'setosa'
```

```python

```

Output:
```
<__main__.LabeledPt at 0x7f6bb73a30d0>
```

<!-- grade_id: cell-c890ae7b76509760, points: 5, grade: true, solution: false -->
```python

```

**Problem**

* Use ```defaultdict``` to organize the points according to species

**Output:** 

* The output should be an object of type ```collections.defaultdict``` called ```species_points``` containing the points organized by species

**Hints:**

```python
from collections import defaultdict
```

Here's an example of how to use a `defaultdict` from the Python docs:

https://docs.python.org/2/library/collections.html#defaultdict-objects

```python

>>> s = [('yellow', 1), ('blue', 2), ('yellow', 3), ('blue', 4), ('red', 1)]
>>> d = defaultdict(list)
>>> for k, v in s:
...     d[k].append(v)
...
>>> d.items()
[('blue', [2, 4]), ('red', [1]), ('yellow', [1, 3])]

```

* Use a similar process to organize the data according to species

```python
from collections import defaultdict

d = defaultdict(list)
for k.label, v.point in labeled_iris_points:
  d[k].append(v)
d.items()
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-111-d46fedc78456> in <module>
      2 
      3 d = defaultdict(list)
----> 4 for k.label, v.point in labeled_iris_points:
      5   d[k].append(v)
      6 d.items()

TypeError: cannot unpack non-iterable LabeledPt object
```

```python

```

```python

```

```python

```

<!-- grade_id: cell-3d34842e5ba7c4c2, points: 0, grade: false, solution: true -->
```python
s = [('yellow', 1), ('blue', 2), ('yellow', 3), ('blue', 4), ('red', 1)]
>>> d = defaultdict(list)
>>> for k, v in s:
...     d[k].append(v)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-106-93c9516e896d> in <module>
      1 s = [('yellow', 1), ('blue', 2), ('yellow', 3), ('blue', 4), ('red', 1)]
----> 2 d = defaultdict(list)
      3 for k, v in s:
      4     d[k].append(v)

NameError: name 'defaultdict' is not defined
```

<!-- grade_id: cell-d69579fedc7abf15, points: 5, grade: true, solution: false -->
```python

```

Run this code to split the data into "Training" data and "Testing" data

```python
import random

def data_splitter(data, prob):
    data = data[:]                    
    random.shuffle(data)              
    cut = int(len(data) * prob)       
    return data[:cut], data[cut:]     

random.seed(12)
iris_X_train, iris_X_test = data_splitter(labeled_iris_points, 0.70)
```

**Problem:**

After creating our KNN classifier, we are interested in seeing how accurately it can classify points according to the measurement values.

Run the KNN classifier with **5 neighbors** on the training dataset, `iris_X_train` and compute the percentage of correct predictions using the testing dataset `iris_X_test`. 

This means that we are labeling the new points in the test data according to the majority vote of the labels of the 5 nearest neighbors in the train data. 

**Output:**

The output should be a float ``pct_correct`` that represents the percentage of samples in the testing dataset that are classified correctly. For example, if there are 10 testing samples, and the classifier labels 9 of the tests correctly, then `pct_correct=90.0`

**Hints:**

* `pct_correct = num_correct / len(iris_X_test) * 100`
* Use a for loop to iterate through the test samples in `iris_X_test`
* You can keep track of the number of correct predictions using a variable

<!-- grade_id: cell-552c5824f73e49be, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-abada96259c3fba4, points: 10, grade: true, solution: false -->
```python

```